```python
# Initialize Otter
import otter
grader = otter.Notebook("hw06.ipynb")
```

<img src="./ccsf.png" alt="CCSF Logo" width=200px style="margin:0px -5px">

# Homework 6: Simulation and Chance

## References

* [Sections 9.0 - 9.3](https://ccsf-math-108.github.io/textbook/chapters/09/Randomness.html)
* [Chapter 10](https://ccsf-math-108.github.io/textbook/chapters/10/Sampling_and_Empirical_Distributions.html)
* [`datascience` Documentation](https://datascience.readthedocs.io/)
* [Code Reference](https://ccsf-math-108.github.io/materials-sp25/resources/code-reference.html)

---

## Assignment Reminders

- 🚨 Make sure to run the code cell at the top of this notebook that starts with `# Initialize Otter` to load the auto-grader.
- Your Tasks are categorized as auto-graded (📍) and manually graded (📍🔎).
    - For all the auto-graded tasks:
        - Replace the `...` in the provided code cell with your own code.
        - Run the `grader.check` code cell to run some tests on your code.
        - Keep in mind that for homework and project assignments, sometimes there are hidden tests that you will not be able to see the results of that we use for scoring the correctness of your response. **Passing the auto-grader does not guarantee that your answer is correct.**
    - For all the manually graded tasks:
        - You might need to provide your own response to the provided prompt. Do so by replacing the template text "_Type your answer here, replacing this text._" with your own words.
        - You might need to produce a graphic or something else using code. Do so by replacing the `...` in the code cell to generate the image, table, etc.
        - In either case, review the rubric on the associated <a href="https://ccsf.instructure.com" target="_blank">Canvas</a> Assignment page to understand the scoring criteria.
- Throughout this assignment and all future ones, please be sure to not re-assign variables throughout the notebook! _For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!_
- You may [submit](#Submit-Your-Assignment-to-Canvas) this assignment as many times as you want before the deadline. Your instructor will score the last version you submit once the deadline has passed.
- We encourage you to discuss this assignment with others but make sure to write and submit your own code. Refer to the syllabus to learn more about how to learn cooperatively.

---

## Configure the Notebook

Run the following cell to configure this Notebook.

```python
from datascience import *
import numpy as np
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
```

---

## Roulette

---

### Chance

A Nevada roulette wheel has 38 pockets and a small ball that rests on the wheel. When the wheel is spun, the ball comes to rest in one of the 38 pockets. That pocket is declared the winner. 

The pockets are labeled 0, 00, 1, 2, 3, 4, ..., 36. Pockets 0 and 00 are green, and the other pockets are alternating red and black. The table `wheel` is a representation of a Nevada roulette wheel. **Note that *both* columns consist of strings.** Below is an example of a roulette wheel!

<img src="./roulette_wheel.jpeg" alt="roulette wheel" width="330px">

Run the cell below to load the `wheel` table.

```python
wheel = Table.read_table('roulette_wheel.csv', dtype=str)
wheel
```

Before you do the following tasks, make sure you understand the logic behind all the examples in [Section 9.5](https://ccsf-math-108.github.io/textbook/chapters/09/5/Finding_Probabilities.html). 

Good ways to approach probability calculations include:

- Thinking one trial at a time: What does the first one have to be? Then what does the next one have to be?
- Breaking up the event into distinct ways in which it can happen. Listing all the distinct ways the event can occur.
- Seeing if it is easier to find the chance that the event does not happen.

On each spin of a roulette wheel, all 38 pockets are equally likely to be the winner regardless of the results of other spins. Among the 38 pockets, 18 are red, 18 black, and 2 green.

---

#### Task 01 📍

_Provide an expression that Python evaluates to the chance of the event described._

Event: The winning pocket is black on all of the first three spins.

_Points:_ 2

```python
first_three_black = ...
```

```python
grader.check("task_01")
```

---

#### Task 02 📍

_Provide an expression that Python evaluates to the chance of the event described._

Event: The color green never wins in the first 10 spins.

_Points:_ 2

```python
no_green = ...
```

```python
grader.check("task_02")
```

---

#### Task 03 📍

_Provide an expression that Python evaluates to the chance of the event described._

Event: The color green wins **at least** once on the first 10 spins.

_Points:_ 2

```python
at_least_one_green = ...
```

```python
grader.check("task_03")
```

---

### Comparing Chances

In each of the following two tasks, two events A and B are described. Answer the question 'Which of the events is more likely?' by choosing from the following three options and setting the variable in the task to the single integer that corresponds to the option you selected.

1. Event A is more likely than Event B
2. Event B is more likely than Event A
3. The two events have the same chance.

You should be able to make the choices **without calculation**. Good ways to approach this exercise include imagining carrying out the chance experiments yourself, one trial at a time.

---

#### Task 04 📍

A child picks four times at random from a box that has four toy animals: a bear, an elephant, a giraffe, and a kangaroo.

- Event A: all four different animals are picked, assuming the child picks without replacement
- Event B: all four different animals are picked, assuming the child picks with replacement

_Points:_ 2

```python
toys_option = ...
```

```python
grader.check("task_04")
```

---

#### Task 05 📍

In a lottery, two numbers are drawn at random without replacement from the integers 1 through 1000.

- Event A: The number 8 is picked on both draws
- Event B: The same number is picked on both draws


_Points:_ 2

```python
lottery_option = ...
```

```python
grader.check("task_05")
```

---

## Simulation with Roulette

---

### Betting on Red

If you bet on *red*, you are betting that the winning pocket will be red. This bet *pays 1 to 1*. That means if you place a one-dollar bet on red, then:

- If the winning pocket is red, you gain 1 dollar. That is, you get your original dollar back, plus one more dollar.
- if the winning pocket is not red, you lose your dollar. In other words, you gain -1 dollars.

Let's see if you can make money by betting on red at roulette.

---

#### Task 06 📍

Define a function `dollar_bet_on_red`. The function definition should:
1. Have one argument `color` that is a `str` for the name of a color.
2. Return your gain in dollars as an `int` if that color had won and you had placed a one-dollar bet on red. 

Consider the following as you work:
* Remember that the gain can be negative.
* Make sure your function returns an integer.
* You can assume that the only colors that will be passed as arguments are `'red'`, `'black'`, and `'green'`. Your function doesn't have to check that the input is correct.

_Points:_ 3

```python
def dollar_bet_on_red(...):
    ...
```

```python
grader.check("task_06")
```

Run the cell below to make sure your function is working.

```python
print(dollar_bet_on_red('green'))
print(dollar_bet_on_red('black'))
print(dollar_bet_on_red('red'))
```

---

#### Task 07 📍

1. Add a column labeled `'Winnings: Red'` as the last column in the table `wheel` such that, for each pocket, the column should contain your gain in dollars if that pocket won and you had bet one dollar on red. 
2. Your code should apply the function `dollar_bet_on_red` on the `wheel` table to create an array `red_winnings`.

_Points:_ 3

```python
red_winnings = ...
wheel = ...
wheel
```

```python
grader.check("task_07")
```

---

### Simulating 10 Bets on Red

Roulette wheels are set up so that each time they are spun, the winning pocket is equally likely to be any of the 38 pockets regardless of the results of all other spins. Let's see what would happen if we decided to bet one dollar on red each round.

---

#### Task 08 📍

1. Create a table `ten_bets` by sampling the table `wheel` to simulate 10 spins of the roulette wheel. 
2. Your table should have the same three column labels as in `wheel`. 
3. Once you've created that table, set `sum_bets` to your net gain, an integer, in all 10 bets, assuming that you bet one dollar on red each time.

While you are working on this, it may be helpful to print out `ten_bets` after you create it!

_Points:_ 5

```python
ten_bets = ...
sum_bets = ...(ten_bets.column('Winnings: Red'))
sum_bets
```

```python
grader.check("task_08")
```

Run the cells above a few times to see how much money you would make if you made 10 one-dollar bets on red. Making a negative amount of money doesn't feel good, but it is a reality in gambling. Casinos are a business, and they make money when gamblers lose.

---

#### Task 09 📍

Let's see what would happen if you made more bets. 

Define a function `net_gain_red` with one argument that takes the number of bets (an integer) and returns the net gain in that number of one-dollar bets on red (an integer).

Reference the `wheel` table that you've defined in the notebook within your function.

_Points:_ 2

```python
def net_gain_red(...):
    ...
    
net_gain_red(10)
```

```python
grader.check("task_09")
```

---

#### Task 10 📍

1. Complete the cell below to simulate the net gain in 200 one-dollar bets on red, repeating the process 10,000 times. 
2. After the cell is run, `all_gains_red` should be an array with 10,000 entries, each of which is the net gain in 200 one-dollar bets on red.

_Points:_ 2

```python
num_bets = ...
repetitions = ...

all_gains_red = ...
...

len(all_gains_red) # Do not change this line! Check that all_gains_red is length 10000.
```

```python
grader.check("task_10")
```

Run the cell below to visualize the results of your simulation.

```python
gains = Table().with_columns('Net Gain on Red', all_gains_red)
gains.hist(bins = np.arange(-80, 41, 4))
```

---

#### Task 11 📍

Using the histogram above, decide whether the following statement is true or false:

>If you make 200 one-dollar bets on red, your chance of losing money is more than 50%.

Assign `loss_more_than_50` to either `True` or `False` depending on your answer to the question.

_Points:_ 2

```python
loss_more_than_50 = ...
```

```python
grader.check("task_11")
```

---

### Betting on a Split

If betting on red doesn't seem like a good idea, maybe a gambler might want to try a different bet. A bet on a *split* is a bet on two consecutive numbers such as 5 and 6. This bets pays 17 to 1. That means if you place a one-dollar bet on the split 5 and 6, then:

- If the winning pocket is either 5 or 6, your gain is 17 dollars.
- If any other pocket wins, you lose your dollar, so your gain is -1 dollars.

---

#### Task 12 📍

Define a function `dollar_bet_on_5_6_split`.
1. The function should have 1 argument (type `str`) that represents the pocket number
2. The function should return the gain in dollars if that pocket won and they had bet one dollar on the 5-6 split.

Remember that the pockets are represented as strings.

_Points:_ 4

```python
def dollar_bet_on_5_6_split(...):
    ...
```

```python
grader.check("task_12")
```

Run the cell below to check that your function is doing what it should.

```python
print(dollar_bet_on_5_6_split('5'))
print(dollar_bet_on_5_6_split('6'))
print(dollar_bet_on_5_6_split('00'))
print(dollar_bet_on_5_6_split('23'))
```

---

#### Task 13 📍

Add a column `'Winnings: 5-6 Split'` to the end of the `wheel` table. For each pocket, the column should contain your gain in dollars if that pocket won and you had bet one dollar on the 5-6 split.

_Points:_ 3

```python
split_winnings = ...
wheel = ...
```

```python
grader.check("task_13")
```

---

#### Task 14 📍

1. Simulate the net gain in 200 one-dollar bets on the 5-6 split.
2. Repeat the simulation for a total of 10,000 times.
3. Store your gains in the array `all_gains_split`.

_Points:_ 2

```python
all_gains_split = ...
...

# Do not change the two lines below
gains = gains.with_columns('Net Gain on Split', all_gains_split)
gains.hist(bins = np.arange(-200, 150, 20))
```

```python
grader.check("task_14")
```

---

#### Task 15 📍

Look carefully at the visualization above, and assign `histogram_statements` to an array of the numbers of each statement below that can be correctly inferred from the overlaid histogram.

1. If you bet one dollar 200 times on a split, your chance of losing money is more than 50%.
2. If you bet one dollar 200 times in roulette, your chance of making more than 50 dollars is greater if you bet on a split each time than if you bet on red each time.
3. If you bet one dollar 200 times in roulette, your chance of losing more than 50 dollars is greater if you bet on a split each time than if you bet on red each time.

Notice that you've already seen one of these statements in a prior question.

_Points:_ 5

```python
histogram_statements = ...
```

```python
grader.check("task_15")
```

If this exercise has put you off playing roulette, it has done its job. If you are still curious about other bets, [here](https://en.wikipedia.org/wiki/Roulette#Bet_odds_table) they all are, and [here](https://en.wikipedia.org/wiki/Roulette#House_edge) is the bad news. The house – that is, the casino – always has an edge over the gambler.

---

## Submit Your Assignment to Canvas

Follow these steps to submit your homework assignment:

1. **Review the Rubric:** View the rubric on the associated Canvas Assignment page to understand the scoring criteria.
2. **Run the Auto-Grader:** Ensure you have executed the code cell containing the command `grader.check_all()` to run all tests for auto-graded tasks marked with 📍. This command will execute all auto-grader tests sequentially.
3. **Complete Manually Graded Tasks:** Verify that you have responded to all the manually graded tasks marked with 📍🔎.
4. **Save Your Work:** In the notebook's Toolbar, go to `File -> Save Notebook` to save your work and create a checkpoint.
5. **Download the Notebook:** In the notebook's Toolbar, go to `File -> Download IPYNB` to download the notebook (`.ipynb`) file.
6. **Upload to Canvas:** On the Canvas Assignment page, click "Start Assignment" or "New Attempt" to upload the downloaded `.ipynb` file.

---

## Attribution

This content is licensed under the <a href="https://creativecommons.org/licenses/by-nc-sa/4.0/">Creative Commons Attribution-NonCommercial-ShareAlike 4.0 International License (CC BY-NC-SA 4.0)</a> and derived from the <a href="https://www.data8.org/">Data 8: The Foundations of Data Science</a> offered by the University of California, Berkeley.

<img src="./by-nc-sa.png" width=100px>

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for task_01 (points: 2)
```
>>> 0 <= first_three_black <= 1
True
```
failure_message: ❌ first_three_black should be a numerical value between 0 and 1, inclusive.
success_message: ✅ first_three_black is a numerical value between 0 and 1, inclusive.


Tests for task_02 (points: 2)
```
>>> 0 <= no_green <= 1
True
```
failure_message: ❌ no_green should be a numerical value between 0 and 1, inclusive.
success_message: ✅ no_green is a numerical value between 0 and 1, inclusive.


Tests for task_03 (points: 2)
```
>>> 0 <= at_least_one_green <= 1
True
```
failure_message: ❌ at_least_one_green should be a numerical value between 0 and 1, inclusive.
success_message: ✅ at_least_one_green is a numerical value between 0 and 1, inclusive.


Tests for task_04 (points: 2)
```
>>> toys_option in [1, 2, 3]
True
```
failure_message: ❌ 1, 2, or 3 should be assigned to toys_option.
success_message: ✅ This is a possible value for toys_option.


Tests for task_05 (points: 2)
```
>>> lottery_option in [1, 2, 3]
True
```
failure_message: ❌ 1, 2, or 3 should be assigned to lottery_option.
success_message: ✅ This is a possible value for lottery_option.


Tests for task_06 (points: 3)
```
>>> dollar_bet_on_red('red') == 1
True
```
failure_message: ❌ Your function should return 1 when it is called on 'red'.
success_message: ✅ Your function correctly returns 1 when it is called on 'red'.


Tests for task_07 (points: 3)
```
>>> isinstance(red_winnings, np.ndarray)
True
```
failure_message: ❌ red_winnings should be an array.
success_message: ✅ It seems that red_winnings is an array.
```
>>> wheel.labels[2] == 'Winnings: Red'
True
```
failure_message: ❌ wheel is not a table with the correct labels.
success_message: ✅ wheel is a table with the correct labels.


Tests for task_08 (points: 5)
```
>>> ten_bets.num_rows == 10
True
```
failure_message: ❌ ten_bets should have 10 rows.
success_message: ✅ ten_bets has 10 rows.
```
>>> set(['Pocket', 'Color', 'Winnings: Red']) == set(ten_bets.labels)
True
```
failure_message: ❌ ten_bets should have the same column labels as wheels.
success_message: ✅ ten_bets has the same column labels as wheels.
```
>>> isinstance(sum_bets, int) or isinstance(sum_bets, np.int64)
True
```
failure_message: ❌ sum_bets should be an integer.
success_message: ✅ sum_bets is an integer.


Tests for task_09 (points: 2)
```
>>> net_gain_red(10000) != net_gain_red(10000)
True
```
failure_message: ❌ There is essentially no chance that the net gain for 2 rounds of 10,000 bets should the same.
success_message: ✅ There is essentially no chance that the net gain for 2 rounds of 10,000 bets should the same.


Tests for task_10 (points: 2)
```
>>> len(all_gains_red) == 10000
True
```
failure_message: ❌ all_gains_red should be an array with 10,000 items.
success_message: ✅ all_gains_red is an array with 10,000 items.


Tests for task_11 (points: 2)
```
>>> isinstance(loss_more_than_50, bool)
True
```
failure_message: ❌ loss_more_than_50 should be a bool value.
success_message: ✅ loss_more_than_50 is a bool value.


Tests for task_12 (points: 4)
```
>>> dollar_bet_on_5_6_split('5') == 17
True
```
failure_message: ❌ Your function didn't return the correct value for the 5 pocket.
success_message: ✅ Your function returned the correct value for the 5 pocket.
```
>>> dollar_bet_on_5_6_split('00') == -1
True
```
failure_message: ❌ Your function didn't return the correct value for the 00 pocket.
success_message: ✅ Your function returned the correct value for the 00 pocket.


Tests for task_13 (points: 3)
```
>>> wheel.labels[3] == 'Winnings: 5-6 Split'
True
```
failure_message: ❌ Your last column label should be 'Winnings: 5-6 Split'.
success_message: ✅ Your last column label is 'Winnings: 5-6 Split'.
```
>>> wheel.where('Pocket', '5').column('Winnings: 5-6 Split').item(0) == 17
True
```
failure_message: ❌ The added column value associated with the 5 pocket is not correct.
success_message: ✅ The added column value associated with the 5 pocket is correct.


Tests for task_14 (points: 2)
```
>>> len(all_gains_split) == 10000
True
```
failure_message: ❌ There should be 10,000 items in all_gains_split.
success_message: ✅ There are 10,000 items in all_gains_split.


Tests for task_15 (points: 5)
```
>>> isinstance(histogram_statements, np.ndarray)
True
```
failure_message: ❌ histogram_statements should be an array.
success_message: ✅ histogram_statements is an array.
```
>>> isinstance(histogram_statements.item(0), int)
True
```
failure_message: ❌ histogram_statements should contain at least one integer.
success_message: ✅ histogram_statements contains at least one integer.
